---
## Part A (2 points)

Write a function that returns a list of numbers, such that
$$
x_i = i^2 \quad \text{for} \quad 1 \leq i \leq n
$$

You may assume $n$ is always greater than 0.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    """
    ### BEGIN SOLUTION
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
print(squares(10))
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: correct_squares, points: 2, grade: true, solution: false -->
```python
"""
Run this cell to check your answer.
Note: there may be additional tests that are added after you submit your assignment, but this
should give a good indication if your answer is correct or not.
"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
### BEGIN HIDDEN TESTS
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
assert squares(20) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 196, 225, 256, 289, 324, 361, 400]
### END HIDDEN TESTS
```

---

## Part B (2 point)

Using your `squares` function above, write another function that computes the sum of the squares of the numbers from 1 to $n$. Your function must use the `squares` function above -- it must NOT re-implement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    ### BEGIN SOLUTION
    return sum(squares(n))
    ### END SOLUTION
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
print(sum_of_squares(10))
```

Output:
```
385
```

<!-- grade_id: correct_sum_of_squares, points: 2, grade: true, solution: false -->
```python
"""
Run this cell to check your answer.
Note: there may be additional tests that are added after you submit your assignment, but this
should give a good indication if your answer is correct or not.
"""
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(10) == 385
### BEGIN HIDDEN TESTS
assert sum_of_squares(11) == 506
assert sum_of_squares(20) == 2870

# Check that sum_of_squares relies on squares.
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
### END HIDDEN TESTS
```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
$$
\sum_{i=1}^n i^2
$$